Please read the following instructions thoroughly. Neglecting to do so may result in missed points.

### Preamble
**Reminder**: Homeworks are due by 7:00PM ET on Sundays.

Before you turn this problem set in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

### Naming conventions
Be sure the filename of your notebook is in the following form:

    <uni>_<assignment>_<details [optional]>.<extension>
    
For example:

    lr3086_hw01.ipynb
    lr3086_hw01_complete.ipynb
    LR3086_HW01.ipynb
    
To rename a notebook, in the menubar, select File$\rightarrow$Rename. The extension for notebook files, `.ipynb`, will already be appended to the filename, but will be hidden from view within the notebook.
    
This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

Courseworks may rename your file to something like `lr3086_hw0-1.ipynb` if you resubmit your assignment. This is perfectly fine.

### What Format To Submit In

Most homeworks are in Jupyter notebooks. Once you've finished your homework, unless specified otherwise, please download your work as an `.ipynb` file to your local machine, then upload it to Courseworks when complete (in the menubar, select File$\rightarrow$Download as$\rightarrow$Notebook).

**Failure to submit a Jupyter notebook will result in a grade of zero for the assignment.**

### Grading

Possible points on late assignments are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors.

Because grading is automated, please delete (or comment out) the `raise NotImplmeneted` code before attempting a problem.

Empty un-editable cells in an assignment are there for a reason. They will be filled with tests by the automatic grader. Please do not attempt to remove them.

### Getting Help

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. You can not post screenshots of your current work to Ed Discussions or other tools used for getting help.

If you need to reach out to a CA for help, please do so via Ed Discussions and not via email. Answers given via Ed Discussions will help you as well as other students. Thus, emails will always have a lower priority for response than Ed Discussions questions. If you do email the CA, please make a note of what section you are in. This helps us identify you in Courseworks faster.

Finally, if you do not get a repsonse from a CA within 48 hours, you may email the professor.

---

# Homework 4: Object Oriented Programming

Total Questions: 6<br/>
Total Points: 10

## Context

Questions 1 through 5 will have you practice some object oriented programming by implementing an **Address Book** for managing contact information.

You will be asked to implement a number of new types, and some methods on them for implementing address book functionality you likely are already familiar with.

When asked to modify an existing type in an exercise, you should completely redefine the type -- in other words, if you answered:

```py
class Something:
    def __init__(self):
        pass
```

for the first question, and are then asked to add a new `foo` attribute to `Something` instances in question 2, your answer to question 2 should be:

```py
class Something:
    def __init__(self):
        self.foo = 12
```

where you have completely redefined the Something class with the newly needed attribute.

## Question 1

We first will start with a type which represents each **contact** we store in the address book.

We'll want a contact to have a (full) name which is required, zero or more phone numbers, and zero or more email addresses for the contact.

Here's an example of how someone will want to use our class:

```py
karina = Contact(
    name="Karina Whitley", 
    addresses=[], 
    email_addresses=["kwhitley@example.org"]
)
```

Write a `Contact` class supporting the above, which additionally re-exposes the `name`, `addresses`, and `email_addresses` arguments as public attributes. In other words, in the above example, `karina.name` should return `"Karina Whitley"`, the string that was passed in.

Remember that both `addresses` and `email_addresses` support 0 or more addresses. To make providing 0 of these more convenient, your class should support leaving those parameters off, meaning:

```py
tayte = Contact(name="Tayte Whitaker")
```

should create a contact whose `addresses` and `email_addresses` are set to a falsy value (any falsy value is acceptable, like `None` or an empty list, `[]`).

[2 points]

<!-- grade_id: cell-58ef6d0845f50544, points: 0, grade: false, solution: true -->
```python
class Contact:
    def __init__(self, name, addresses=None, email_addresses=None):
        self.name = name
        self.addresses = addresses or []
        self.email_addresses = email_addresses or []
```

<!-- grade_id: cell-4086caa50583bb00, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
karina = Contact(
    name="Karina Whitley", 
    addresses=[], 
    email_addresses=["kwhitley@example.org"]
)
assert "Karina Whitley" == karina.name
assert not karina.addresses
assert ["kwhitley@example.org"] == karina.email_addresses

### END TESTS
```

<!-- grade_id: cell-cff1aad36e1aae6f, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
edwin = Contact(
    name="Edwin Alder", 
    addresses=["100 Park Avenue South\nNew York, NY 10003"], 
    email_addresses=["ealder@example.org"]
)

assert (edwin.name, edwin.addresses, edwin.email_addresses) == (
    "Edwin Alder",
    ["100 Park Avenue South\nNew York, NY 10003"],
    ["ealder@example.org"],
)
### END TESTS
```

<!-- grade_id: cell-ec15216391d25859, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
tae = Contact(name="Tae Danae Ruby")
assert tae.name == "Tae Danae Ruby"
assert not tae.addresses
assert not tae.email_addresses
### END TESTS
```

<!-- grade_id: cell-bd391629be5f0670, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
christine = Contact(
    name="Christine Robins", 
    email_addresses=["crobins@example.com", "crobins@example.edu"]
)
assert christine.name == "Christine Robins"
assert not christine.addresses
assert christine.email_addresses == ["crobins@example.com", "crobins@example.edu"]
### END TESTS
```

## Question 2

Now that we have a type which represents an individual contact (person), we want a type which represents our address book which will contain lots of contacts.

Write an `AddressBook` class which takes a list of contacts (defined as above), and can be used as such:

```py
lynn = Contact(name="Lynn Root")
book = AddressBook(contacts=[lynn])
```

Give your address book class a _method_ called `add_contact` which can add a contact to the address book even after creating one.

In other words:

```py
sai = Contact(
    name="Sai Lee", 
    email_addresses=["sai@example.com"]
)
book.add_contact(sai)
```
which should add the new contact to the address book.

[2 points]

<!-- grade_id: cell-6c24414d136a3f22, points: 0, grade: false, solution: true -->
```python
class AddressBook:
    def __init__(self, contacts=[]):
        self.contacts = contacts

    def add_contact(self, contact):
        self.contacts.append(contact)
```

<!-- grade_id: cell-3ac167f97aea38ca, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
lynn = Contact(name="Lynn Root")
book = AddressBook(contacts=[lynn])
book.add_contact(Contact(name="Sai Lee", email_addresses=["sai@example.com"]))
### END TESTS
```

<!-- grade_id: cell-4040a0858c2c2d70, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
book = AddressBook(contacts=[])
book.add_contact(lynn)
### END TESTS
```

## Question 3

Our address book can have contacts added to it, but we have no way of looking up their information.

Let's add one.

Redefine the `AddressBook` class such that it also has a `get_contact_by_name` method which takes a name (a `str`) and returns the contact with the given name -- i.e. it returns an instance of `Contact`.

You can assume that all names are unique, so no two contacts will have the same exact name for this exercise.

Implement `get_contact_by_name` method so that you only have to provide the string of the contact's name and not a `Contact` instance.

Be sure to also include the implementation for the `add_contact` method you defined from question 2.

[2 points]

<!-- grade_id: cell-762a12f6641e3f64, points: 0, grade: false, solution: true -->
```python
class AddressBook:
    def __init__(self, contacts=[]):
        self.contacts = contacts

    def add_contact(self, contact):
        self.contacts.append(contact)

    def get_contact_by_name(self, name):
        for contact in self.contacts:
            if contact.name == name:
                return contact
        return None
```

<!-- grade_id: cell-a526d1180a86e905, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
lynn = Contact(name="Lynn Root")
book = AddressBook(contacts=[lynn])
assert book.get_contact_by_name("Lynn Root") == lynn
### END TESTS
```

<!-- grade_id: cell-8a64b693f8bafd7e, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
sai = Contact(name="Sai Lee", email_addresses=["sai@example.com"])
book.add_contact(sai)
assert book.get_contact_by_name("Lynn Root") == lynn
assert book.get_contact_by_name("Sai Lee") == sai
### END TESTS
```

## Question 4

We can now both add contacts to the address book as well as look them up.

But we want to be able to signal that a contact doesn't exist if we attempt to look up one by name that is not present in the address book.

Define an exception called `NoSuchContact`, and raise this exception in `get_contact_by_name` if it is called with a name that does not match any contact. You should include the name of the (non-existing) contact in the exception message for `NoSuchContact`. 

Don't forget, you'll have to re-define the `AddressBook` class in order to contain this new logic.

[2 points]

<!-- grade_id: cell-e86904ff06b67d34, points: 0, grade: false, solution: true -->
```python
class NoSuchContact(Exception):
    pass

class AddressBook:
    def __init__(self, contacts=[]):
        self.contacts = contacts

    def add_contact(self, contact):
        self.contacts.append(contact)

    def get_contact_by_name(self, name):
        for contact in self.contacts:
            if contact.name == name:
                return contact
        raise NoSuchContact(f"No contact found with the name '{name}'")
 
```

<!-- grade_id: cell-a98e218fe0d1fd39, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
lynn = Contact(name="Lynn Root")
book = AddressBook(contacts=[lynn])
assert book.get_contact_by_name("Lynn Root") == lynn
try:
    book.get_contact_by_name("Unknown Contact")
except NoSuchContact:
    pass
else:
    assert False, "Expected a 'NoSuchContact' exception!"
### END TESTS
```

<!-- grade_id: cell-a153e8f8472e06bf, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
lynn = Contact(name="Lynn Root")
book = AddressBook(contacts=[lynn])
assert book.get_contact_by_name("Lynn Root") == lynn
try:
    book.get_contact_by_name("Unknown Contact")
except NoSuchContact as e:
    assert "Unknown Contact" in str(e)
else:
    assert False, "Expected a 'NoSuchContact' exception!"
### END TESTS
```

<!-- grade_id: cell-5de2fe39e3024da1, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
try:
    book.get_contact_by_name("Another Unknown Contact")
except NoSuchContact as e:
    assert "Another Unknown Contact" in str(e)
else:
    assert False, "Expected a 'NoSuchContact' exception!"
### END TESTS
```

## Question 5

Write an implementation of the `__repr__` special method which for a `Contact` shows the name of the contact.

For a contact named `"Sara Loggins"`, we want to see `"Contact(name='Sara Loggins')"`. You should not include any other contact fields in the output. Be sure that your repr output exactly matches the example shown!

[1 point]

<!-- grade_id: cell-f319c54ef28284b0, points: 0, grade: false, solution: true -->
```python
class Contact:
    def __init__(self, name, addresses=None, email_addresses=None):
        self.name = name
        self.addresses = addresses or []
        self.email_addresses = email_addresses or []

    def __repr__(self):
        return f"Contact(name='{self.name}')"
```

<!-- grade_id: cell-8682a765f98c409d, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert repr(Contact(name="Gandalf Hernandez")) == "Contact(name='Gandalf Hernandez')"
### END TESTS
```

<!-- grade_id: cell-1950365786765918, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert repr(Contact(name="Lynn Root")) == "Contact(name='Lynn Root')"
### END TESTS
```

## Question 6

Create a class named `String` that inherits from `str`. Add a method to your subclass that checks if the string does not start with a given string. Call that method `notstartswith`. This class should only take up 3 lines of code however you will not be graded on the length of your class.

[1 point]

<!-- grade_id: cell-63e0e838fde67499, points: 0, grade: false, solution: true -->
```python
class String(str):
    def notstartswith(self, prefix):
        return not self.startswith(prefix)
```

<!-- grade_id: cell-437707ddadd1642a, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
what_does_the_dog_say = "woof"
word = String(what_does_the_dog_say)

assert word.notstartswith("w") is False
### END TESTS
```

<!-- grade_id: cell-3fb8e83ec56679c5, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK
```

<!-- grade_id: cell-b59882bc0ba6b144, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK
```